# Welcome to Lab: Probability 🎂

### Happy Birthday!

Just kidding, it's *(probably, with P=0.9973)* not your birthday, but it might be someone's birthday in this room! 

In this lab, we will be completing some classic probability questions surrounding birthdays!

- For these puzzles, we assume that each year is **365 days** (ignore leap years)
- We also assume that everyone has an **equal chance** of a birthday each day *(not actually true in real life)*


A few tips to remember:

- **You are not alone on your journey in learning programming!**  You have your lab Teaching Assistant, your Course Aides, your lab group, and the professors (Prof. Wade and Prof. Karle), who are all here to help you out!
- If you find yourself stuck for more than a few minutes, ask a neighbor or course staff for help!  When you are giving help to your neighbor, explain the **idea and approach** to the problem without sharing the answer itself so they can have the same **<i>ah-hah</i>** moment!
- We are here to help you!  Don't feel embarrassed or shy to ask us for help!

Let's get started!

```python
# Meet your CAs and TA if you haven't already!
# ...first name is enough, we'll know who they are! :)
ta_name = ""
ca1_name = ""
ca2_name = ""
ca3_name = ""


# Say hello to each other!
# - Groups of 3 are ideal :)
# - However, groups of 2 or 4 are fine too!
#
# QOTD to Ask Your Group: "What's your favorite candy?"
partner1_name = ""
partner1_netid = ""
partner1_favcandy = ""

partner2_name = ""
partner2_netid = ""
partner2_favcandy = ""

partner3_name = ""
partner3_netid = ""
partner3_favcandy = ""
```

<hr style="color: #DD3403;">

# Part 1: Your Birthday

The **birthday problem** is a classic statistical problem that asks: *How likely is it for at least two people in a room of **n** people to share a birthday?*  In other words, what's the chance that there's **at least one match of birthdays** in a room with **n** people? We'll eventually answer this question in this lab.

To begin, let's consider **your birthday**.  This is a related, but very different question: 

> *How likely is it for at least one person in a room to have the same birthday as **me**?*

### Puzzle 1.1: Sharing a Birthday

We should first answer: if you are in a room with **one other person**, what is the probability you both have the same birthday? (*Of course, you know your own birthday!*) 

Using **mathematics and probability** (not any special Python functions), replace the ... with your answer in the following puzzles!

First, find the probability that you share the same birthday with **one other person**, storing your numerical answer in `p_one_person_my_birthday`!

*Hint:* How many days of the year make it possible for someone to share **your birthday**?

<!-- grade_id: cell-16021a893dda10f8, points: 0, grade: false, solution: true -->
```python
p_one_person_my_birthday = ...
p_one_person_my_birthday
```

Now, if you are in a room with **two other people**, what is the probability that **at least one of them** has the same birthday as you?

Remember:
- **$P(\text{at least one}) = 1 - P(\text{none})$**
- If someone **does not** share your birthday, that means their birthday is on one of the **other 364 days** of the year. 
- Two people independently have the **same probability** of not sharing a birthday with you, but you would need to **multiply** those probabilities together.
- Store your numerical answer in `p_two_others_my_birthday`!

<!-- grade_id: cell-d776448804b67dfc, points: 0, grade: false, solution: true -->
```python
p_two_others_my_birthday = ...
p_two_others_my_birthday
```

Finally, if you are in a room with **three other people**, what is the probability that **at least one of them** has the same birthday as you?

Store your numerical answer in `p_three_others_my_birthday`. 

Look back to how you calculated `p_two_others_my_birthday` if you're stuck - the calculations involved should be quite similar! 


<!-- grade_id: cell-9b6271534fe56ab1, points: 0, grade: false, solution: true -->
```python
p_three_others_my_birthday = ...
p_three_others_my_birthday
```

### 🔬 Test Case Checkpoint 🔬

<!-- grade_id: cell-2213109ad4022435, points: 0, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 1.1 ==
# - This read-only cell contains test cases for your previous cells.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.
import math 
assert( math.isclose(p_one_person_my_birthday, 0.0027397260273972603) ), "The probability for one birthday doesn't seem right."
assert( math.isclose(p_two_others_my_birthday, 0.005471945956089352) ), "The probability for two birthdays doesn't seem right."
assert( math.isclose(p_three_others_my_birthday, 0.008196680350730179) ), "The probability for three birthdays doesn't seem right."

## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
```

### Puzzle 1.2: Sharing My Birthday with `n` Other People

You may have noticed a pattern beginning to form. When we add more people ($n$) to a room, the probability that **at least one** of them shares a birthday with yours **compounds** (that is, the $P(\text{none})$ in our $P(\text{at least one})$ is multiplied by itself).

Suppose we have a room with **`n` other people**.  

Complete the function `p_my_birthday(n)`, where `n` is the parameter, that **returns the numerical probability** that **at least one person** has the same birthday as you.  All you have to do is replace the ... with the probability written in terms of `n`.

**Note:** A function is a useful way to **reuse code** and use different inputs to get different outputs, similar to a math function except for code! :) Below, we set up most of the code to **define** the function `p_my_birthday` with the parameter `n` as the input (i.e. how many other people there are in the room with you). The output of the function should be the calculated probability of observing **at least one** match with you in a room with `n` other people.


<!-- grade_id: cell-9fcc75505fada2fa, points: 0, grade: false, solution: true -->
```python
# Defining the function
def p_my_birthday(n):
    return ...
```

### Running Your Function

Here are a few inputs to test our function.  Make sure the answers are what you expect. 

You **do not change to do anything** for this cell, **just run it**!

```python
print('A room with one other person:', p_my_birthday(1))
print('A room with two other people:', p_my_birthday(2))
print('A room with three other people:', p_my_birthday(3))
print('A room with 100 other people:', p_my_birthday(100))
print('A room with 1000 other people:', p_my_birthday(1000))
print('A room with 10000 other people:', p_my_birthday(10000))
```

### 🔬 Test Case Checkpoint 🔬

<!-- grade_id: cell-f1a2b09260950b38, points: 0, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 1.2 ==
# - This read-only cell contains test cases for your previous cells.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

assert( math.isclose(p_my_birthday(1), 0.002739726027397249) ), "The probability for your function with n = 1 doesn't seem right."
assert( math.isclose(p_my_birthday(2), 0.005471945956089352) ), "The probability for your function with n = 2 doesn't seem right."
assert( math.isclose(p_my_birthday(3), 0.008196680350730179) ), "The probability for your function with n = 3 doesn't seem right."
assert( math.isclose(p_my_birthday(100), 0.23993292618409912) ), "The probability for your function with n = 100 doesn't seem right."
assert( math.isclose(p_my_birthday(1000), 0.9356543499651447) ), "The probability for your function with n = 1000 doesn't seem right."

## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
```

### Plotting Your Function

The following code below uses your `p_my_birthday` function to **generate a plot** showing the probability of you having the same birthday as someone else **increasing** as the number of people (n) in the room grows.

This code is **complete**, but notice how it is **very close** to the simulation code you've seen in lecture:
- Both code for simulation and the code below starts with an empty list (`data = []`)
- Both code for simulation and the code below runs a defined number of times (`for i in range(2000):`)
- Instead of simulating real-world variables, we record the data for every value of `i`.
- The rest is the same! :)

**Do not change** anything in the following two cells, just **run them and observe**. 

```python
# Record the data in a DataFrame:
import pandas as pd
data = []
for i in range(2000):
    d = { 'Other People': i, 'P(at least one person shares my birthday)': p_my_birthday(i) }
    data.append(d)

df = pd.DataFrame(data)
```

```python
# Plot our DataFrame:
plot = df.plot.line(y='P(at least one person shares my birthday)', x='Other People', grid=True)
```

### Analysis: Birthday Probabilities

**Q1: Is it more or less likely that someone in your lab shares *your* birthday **OR** that there's at least one shared birthday between *anyone* in your lab?**

*(✏️ Edit this cell to replace this text with your answer. ✏️)*

**Group Discussion: In your lab section there is likely over a 50% chance that two people share a birthday. Using the results from Puzzle 1.2 and your conclusions above, discuss with your group whether you think this is intuitive or counterintuitive.**

<hr style="color: #DD3403;">

# Part 2: The Classic Birthday Problem

Now that we've warmed up by considering **your birthday**, we can go back to the original **birthday problem**. 

To restate, the **birthday problem** is a classic problem that asks:
> *How likely is it for at least two people in a room to share a birthday?*

**NOTE**: Since we're no longer using **your** birthday, all birthdays here are now **random**.

### Puzzle 2.1: Birthday Sharing at a Small Party

Let's first walk  through scenarios of people sharing birthdays at super small parties of one, two, or three people. 


#### **A One-Person Party** (n=1)

Suppose there is just one person in a room. Not really a party, but the probability that they **share a birthday** with someone else is... well... 0%. :(

$P(\text{shared birthday with 1 person}) = 0$


#### **A Two-Person Party** (n=2)

Now, let's add one more person to the party. 

The probability of **two people in a room** sharing the same birthday is non-zero! 

This is the probability of the **second person** having **exactly the same birthday** as the first, which is 1 out of 365!

$P(\text{shared birthday with 2 people}) = \frac{1}{365}$


#### **A Three-Person Party** (n=3)

Now it gets interesting! Adding one more person to the party, we now have three total in attendance. 

In a three person party, what is the **probability** that there is **at least one** match of birthdays?

Remember:
- **$P(\text{at least one}) = 1 - P(\text{none})$**
- We aren't only considering **your** birthday, but now **any** birthday in the room of three. 
- For **none of the birthdays** to match, the second person would need to have a birthday on one of the other **364 days**, but the third person would need to have a birthday that **isn't shared** with the unique **first OR second person's birthday**. How many days does that leave? 


Using these hints, find the probability that at least two people share a birthday in a party of three, storing your numerical answer in `p_three_person_party`:

<!-- grade_id: cell-9db7fd67d49459bd, points: 0, grade: false, solution: true -->
```python
p_three_person_party = ...
p_three_person_party
```

### 🔬 Test Case Checkpoint 🔬

<!-- grade_id: cell-3aa5b5a4afef9fbc, points: 0, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 2.1 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.
import math
assert( math.isclose(p_three_person_party, 0.008204165884781345) ), "The probability doesn't seem right."

## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
```

### Puzzle 2.2: A Four-Person Party

Using what you've learned, let's increase our party size to four.

In a four person party, what is the probability that there is **at least one match** of birthdays?

Use the following cell to find the numerical answer, storing it in the variable `p_four_person_party`:

<!-- grade_id: cell-1a5b32da97b97164, points: 0, grade: false, solution: true -->
```python
p_four_person_party = ...
p_four_person_party
```

### 🔬 Test Case Checkpoint 🔬

<!-- grade_id: cell-571231fcdfa9de08, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 2.2 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.
import math
assert( math.isclose(p_four_person_party, 0.016355912466550215) ), "The probability doesn't seem right."

## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
```

### Example: Sharing a Birthday at an `n`-Person Party

For your convenience, we've created a function called `P_shared_birthday` that returns the **probability** that there is **at least one** match of birthdays among `n` people. This function just computes the Birthday Problem we did in lecture a couple weeks ago!

Alongside the function itself, we have also provided **print statements** that will output the **probabilities** of **at least one** match of birthdays among varying `n` people. 

**Do not modify** this function, just **run the cell** below and confirm the **outputs make sense**:

<!-- grade_id: cell-8826e853f191d14f, points: 0, grade: false, solution: true -->
```python
def P_shared_birthday(n):
    p_none = 1
    for i in range(1, n):
        p_none = p_none * (365 - i) / 365
    return 1 - p_none

# A few test outputs...
print('A room with just you:',  P_shared_birthday(1))
print('A room with two people:',  P_shared_birthday(2))
print('A room with three people:',  P_shared_birthday(3))
print('A room with four people:',  P_shared_birthday(4))
print('A room with 40 other people:',  P_shared_birthday(40))
print('A room with 80 other people:',  P_shared_birthday(80))
print('A room with 400 other people:',  P_shared_birthday(400))
```

### 🔬 Test Case Checkpoint 🔬

<!-- grade_id: cell-5a323ce5fab0c281, points: 0, grade: true, solution: false -->
```python
## == TEST CASES for Example Function ==
# - This read-only cell contains test cases for the function we've provided you.
# - Please see a TA/CA if these test cases do not pass. You should not have to change the code we've provided above. 

import math
assert( math.isclose(P_shared_birthday(3), 0.008204165884781456) ), "The P_shared_birthday function is incorrect for n = 3."
assert( math.isclose(P_shared_birthday(4), 0.016355912466550326) ), "The P_shared_birthday function is incorrect for n = 4."
assert( math.isclose(P_shared_birthday(40), 0.891231809817949) ), "The P_shared_birthday function is incorrect for n = 40."
assert( math.isclose(P_shared_birthday(400), 1.0) ), "The P_shared_birthday function is incorrect for n = 400."

## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
```

### Puzzle 2.3: Plotting P_shared_birthday

Now generate a plot showing the growth of `P_shared_birthday` over changes to `n`. For now, use **n = 2000 people**. This plot will display how the **probability of at least two people** sharing a birthday in a room changes as the **number of people grows**.

Your code should be **very similar** to simulation code from lecture.   

If you get stuck, **look back** at how we generated the function plot in **Part 1** of this lab. 



<!-- grade_id: cell-c46ee253ad4a8777, points: 0, grade: false, solution: true -->
```python
# Create the DataFrame:
...
```

```python
# Plot the DataFrame:
...
```

Notice how most of the graph is redundant after a certain point and only a small portion of the graph captures a change in probability as the number of people grows.

Copy and paste your code above into the boxes below, **adjusting the number of people** you record (range of the `for` loop) to focus only on the **interesting part** of this plot (when the probability is changing the most).

```python
# Create the new DataFrame:
...
```

```python
# Plot the new DataFrame:
...
```

### Analysis: Birthday Problem Plot

**Q2: Observe the plots you've generated above. What does the shape of the plots tell you about the probability of at least two people sharing birthdays in a room of `n` people as opposed to them sharing a specific birthday with you? Which one becomes more likely with a lower value of `n` (people in the room)?**

*(✏️ Edit this cell to replace this text with your answer. ✏️)*

<hr style="color: #DD3403;">

# Part 3: Taylor's Birthday Problem

We've made some great progress answering the classic Birthday Problem, but now Taylor Swift is coming to perform at UIUC and wants more. Specifically,  Taylor Swift wants **every day to be someone's birthday**!  

In this portion of the lab, we will discover:
> Given a massive concert of `n` people, what is the probability that **all 365 days have at least one person with a birthday that day**?

Since Wade, Karle, and many of the course staff love Taylor, we wrote a **function** to do this simulation for Taylor!  You can check out the code below, but all you need to know is that:

> The `simulation_for_taylor` function returns **one simulation** of the **number of people** needed to **satisfy Taylor's wish** for there to be someone with a birthday for **every day of the year** at her concert!

...you will need to **use this function** to help us calculate probabilities!

**Do not change** anything in the following cell, just **run it and observe**. 

```python
# The `simulation_for_taylor` function returns one simulation of the number of people needed
# to satisfy Taylor's wish for there to be someone with a birthday for every day of the year
# at her concert!
#
# @return: One simulation of the number of people needed at the concert.
#   Examples: 2056, 2520, 2868, 3033, 2022, 1743, ...
import random

def simulation_for_taylor():
    ct = 0   #< Count the number of people needed
    day = 0  #< Count the current unique days found
    
    # While all 365 unique days have not been seen:
    while day < 365:
        # Add a person:
        ct = ct + 1
        
        # Check if a random number is a new, unique day.
        # - When day == 0, we've seen no unique days so P(unique day) == 1 as randint(0, 364) will always be `>= 0` (100%)
        # - When day == 1, P(unique day) == (364/365) and we check if we get a number `>= 1` (any number but 0)
        # - When day == 2, P(unique day) == (363/365) and we check if we get a number `>= 2` (anything but 0 or 1)
        # - ...
        # - When day == 364, P(unique day) == (1/365) so we need get randint(0, 364) to be exactly 364.
        if random.randint(0, 364) >= day:
            day = day + 1
            
    # Return the number of people needed
    return ct    
```

### Running the Simulation

**Run** the following cell **several times**.  It's calling our `simulation_for_taylor()` function once and outputting a **number**. 

This **number** is the number of people in **one simulated instance** of Taylor's concert **needed** for **every day of the year** to be someone's birthday:

<!-- grade_id: cell-a60a84fa1d64c87d, points: 0, grade: false, solution: true -->
```python
simulation_for_taylor()
```

### Puzzle 3.1: Collecting Simulation Results

Using `simulation_for_taylor()`, collect the results of running this simulation at least **1,000** times and store the result in `df_taylor`. 

Store your results in the DataFrame `df_taylor` under any column name that fits the context. For context, we are populating `df_taylor` with 1,000 simulations that contains the number of people needed to **satisfy Taylor's birthday problem**. 

If you're feeling stuck, remember how we wrote **simulation code** in lecture and **Parts 1 and 2** of this lab.

*(Note: This may take a few seconds or up to a minute to run on a slower laptop.  You may want to try running the simulation 10 times to make sure your code works before changing your `for` loop range to 1,000)*

<!-- grade_id: cell-6a44b9f35daf8325, points: 0, grade: false, solution: true -->
```python
# Simulation code using the `simulation_for_taylor` function to run 1,000 times:
data = []
...

# Finally, create a DataFrame of simulation results:
df_taylor = ...
```

## Puzzle 3.2: Finding the Probability for Taylor

The concert is being held and we have **exactly 2,000 people packed onto the main quad for Taylor Swift!**  

Using your simulation results, we are going to find an **estimate for the probability** that Taylor will be happy to know **every day of the year** is someone’s birthday.

First, using **conditional logic**, find the rows of our `df_taylor` where the **simulation was a success**, storing the results in a DataFrame `df_success`.

*Hint: The simulation is a success if <= 2000 people was enough for Taylor's request.*

<!-- grade_id: cell-0b48a78d683ce577, points: 0, grade: false, solution: true -->
```python
df_success = ...
df_success
```

Now, using your `df_success` DataFrame, find the **probability of success** given all the simulated concerts in our `df_taylor`.

Store this numerical result in the variable `P_every_birthday`. 

<!-- grade_id: cell-3cd0d620430a89ba, points: 0, grade: false, solution: true -->
```python
P_every_birthday = ...
P_every_birthday
```

### Puzzle 3.3: Visualizing Results

With the cell below, generate a histogram of your `df_taylor` results, showing the distribution of people needed in our simulation to account for every birthday.

<!-- grade_id: cell-62147f75639a94d3, points: 0, grade: false, solution: true -->
```python
...
```

### 🔬 Test Case Checkpoint 🔬

<!-- grade_id: cell-1bae19913f8df080, points: 0, grade: true, solution: false -->
```python
## == TEST CASES for Part 3 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

assert( len(df_taylor) >= 1000 ), "Ensure you have done at least 1,000 simulations."
assert( len(df_success) > 0 ), "Ensure your success includes all simulations with 2,000 or fewer people."
assert( len(df_success) < len(df_taylor) ), "Ensure your success includes all simulations with 2,000 or fewer people."
assert( P_every_birthday > 0.1 and P_every_birthday < 0.4 ), "Double check your `P_every_birthday` is a reasonable probability."

## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
```

### Analysis: Taylor's Birthday Problem

**Q3: Looking back at your calculations and histogram of simulation results, comment on whether you think it's likely that Taylor will have her birthday request fulfilled at the UIUC Quad Concert of 2,000 people.**

*(✏️ Edit this cell to replace this text with your answer. ✏️)*

<hr style="color: #DD3403;">

# Submission

You're almost done!  All you need to do is to commit your lab to GitHub:

1.  ⚠️ **Make certain to save your work.** ⚠️ To do this, go to **File => Save All**

2.  After you have saved, exit this notebook and follow the Canvas instructions to commit this lab to your Git repository!

3. Your TA will grade your submission and provide you feedback after the lab is due. :)